Comparative Brain-Wide Mapping of Isoflurane and Ketamine-Activated Nuclei and Functional Networks

Author Response

The following is the authors’ response to the original reviews.

We sincerely thank the editor and reviewers for their constructive feedback on our manuscript. Based on their recommendations, we've conducted additional experiments, made revisions to the text and figures, and provide a point-by-point response below.

Reviewer #1 (Recommendations for the authors):

1. The lack of behavioral/physiological measures of the depth of anesthesia (ventilation, heart rate, blood pressure, temperature, O2, pain reflexes, etc...) combined with the lack of dose-response and the use of different routes of administration makes the data difficult to interpret. Sure, there is a clear difference in network activation between KET and ISO, but are those effects due to the depth of the anesthesia, the route of administration, and the dose used? The lack of behavioral/physiological measures prevents the identification of brain regions responsible for some of the physiological effects and different effects of anesthetics.

We greatly appreciate the insightful feedback you have provided.

In response to the concerns about anesthesia depth:

a. We recorded EEG and EMG data both before and after drug administration. Supplementary Figure 1 showcases the changes in EEG and EMG power observed 30 minutes post-drug administration, normalized to a 5-minute baseline taken prior to the drug's administration. Notably, no significant differences were detected in the normalized EEG and EMG power between the ISO and KET groups. Given the marked statistical differences observed between the EEG power in the KET and saline groups, and the EMG power in the home cage and ISO groups, we infer that both anesthetics effectively induced a loss of consciousness.

b. We used standard methods and doses for inducing c-Fos expression with anesthetics, as documented in prior studies (Hua, T, et al., Nat Neurosci, 2020; 23(7): 854-868; Jiang-Xie, L F, et al., Neuron, 2019; 102(5): 1053-1065.e4; Lu, J, et al., J Comp Neurol, 2008; 508(4): 648-62). In future research, it might be more optimal to adopt continuous intraperitoneal or intravenous administration of ketamine.

c. Within the scope of our study, while disparities in anesthesia duration might potentially influence the direct statistical comparison of ISO and KET, such disparities wouldn't compromise the identification of brain regions activated by KET or ISO when assessed as distinct stimuli (ISO vs. home cage; KET vs. saline) or in relation to their individual functional network hub node results.

We hope these additions and clarifications adequately address your concerns and enhance the comprehensibility of our data.

1. Under anesthesia there should be an overall reduction of activity, is that the case? There is no mention of significantly downregulated regions. The authors use multiple transformations of the data to interpret the results (%, PC1 values, logarithm) without much explanation or showing the full raw data in Fig 1. It would be helpful to interpret the data to compare the average fos+ neurons in each region between treatment and control for each drug.

Absence of Significantly Downregulated Regions Under Anesthesia: There are two primary reasons for this observation:

a. Our study's sampling time for the home cage, ISO, saline, and KET groups was during Zeitgeber Time (ZT) 6-7.5. During this period, mice in both the home cage and saline groups typically showed reduced spontaneous activity or were in a sleep state. Our Supplementary Figure 1 EEG and EMG data corroborate this, revealing no significant statistical variations in EEG power between the home cage and ISO groups, nor in EMG power between the saline and KET groups.

b. Our immunohistochemical data showed that the total number of c-Fos positive cells in the two control groups was notably lower than in the experimental groups (Saline group vs KET group: 11808±2386 versus 308705±106131, P = 0.006; Home cage vs ISO group: 3371±840 vs 12326±1879, P = 0.001). This is in line with previous studies, like the one by Cirelli C and team, which found minimal c-Fos expression throughout the mouse brain during physiological sleep (Cirelli, C, and G Tononi, Sleep, 2000; 23(4): 453-69). Thus, in our analysis, we did not detect regions with significant downregulation when comparing anesthetized mice with controls.

Interpreting Raw Data from Figure 1: Regarding the average Fos+ neurons:

In Figures 4 and 5, we utilized raw data (c-Fos cell count) to assess cell expression differences across 201 brain regions within each group. Only brain regions that had significant statistical differences after multiple comparison corrections are shown in the figures.

1. I do not understand their interpretation of the PCA analyses. For instance, in Fig 2 they claim that KET is associated with PC1 while ISO is associated with PC2. Looking at the distribution of points it's clear that the KET animals are all grouped at around +2.5 on PC1 and -2.0 on PC2, this means that KET is associated with both PC1 and PC2 to a similar degree (2 to 2.5). Moreover, I'm confused about why they use PCA to represent the animals/group. PCA is a powerful technique to reduce dimensionality and identify groups of variables that may represent the same underlying construct; however, it is not the best way to identify clusters of individuals or groups.

Clarification on PCA Analyses in Figure 2: Thank you for pointing out the ambiguities in our initial presentation of the PCA analyses. We are grateful for the opportunity to address these concerns.

KET and ISO Associations with PC1 and PC2: You rightly observed that KET samples manifest both a positive value on PC1 (around +2.5) and a negative one on PC2 (around -2.0), suggesting that KET has a substantial influence on both principal components. In PCA, a positive score implies a positive association with that component, whereas a negative score suggests a negative association. Contrarily, ISO samples predominantly exhibit values around +2.5 on PC2, with nearly neutral values for PC1, underlining its stronger association with PC2 and lack of significant correlation with PC1. To ensure transparency and clarity, we've adjusted the corresponding descriptions in our manuscript, which can be found on Line 100.

Rationale Behind Using PCA to Represent Animals/Groups: Our initial step was to conduct PCA clustering analysis on the 201 brain regions within both the ISO and KET groups. In the accompanying chart, varying colors denote different brain regions, while distinct shapes represent separate clusters. There wasn't a pronounced distribution pattern within the ISO and KET groups, which led us to adopt the current computational method presented in the paper. This approach was chosen to directly contrast the relative differential expressions between ISO and KET.

We deeply value your feedback, which has steered us toward a clearer and more accurate presentation of our data. We genuinely appreciate your meticulous review.

1. The actual metric used for the first PCA is unclear, is it the FOS density in each of the regions (some of those regions are large and consist of many subregions, how does that affect the analysis) is it the %-fos, or normalized cells? The wording describing this is variable causing some confusion. How would looking at these different metrics influence the analysis?

Thank you for raising concerns about the metrics used in our PCA analysis. We recognize the need for clearer exposition and appreciate the opportunity to clarify.

PCA Metrics: The metric for our PCA is calculated by obtaining the ratio of the Fos density within a specific brain region to the global Fos density across the brain. Briefly, this entails dividing the number of Fos-positive cells in a given region by its volume, and then comparing this to the Fos density of the whole brain. The logarithm of this ratio provides our PCA metric. We've elaborated on this in the Materials and Methods section (Lines 401) and enhanced clarity in our revised manuscript, particularly at Line 96.

In Figure 2A, we employed 53 larger, mutually exclusive brain regions based on the reference from the study by Do et al. (eLife, 2016;5:e13214). However, in Figure 3A, we used a more detailed segmentation, incorporating 201 distinct brain areas that are more granular than those in Figure 2A. Notably, the PCA results from both representations were consistent. The rationale behind selecting either the 53 or 201 brain regions can be found in our response to Question 10.

Rationale for Metric Choice: The log ratio of regional c-Fos densities relative to the global brain density was chosen due to:

a. Notable disparities in c-Fos cell expression across the groups.

b. A significant non-normal distribution of density values across animals within the group. Employing the log ratio effectively mitigates the impact of extreme values and outliers, achieving a more standardized data distribution.

We've added PCA plots based on c-Fos densities, depicted in Author response image 2. However, the data dispersion has resulted in a significantly spread-out horizontal scale for these visuals.

1. Based on Fig 3 the authors concludes that ISO activates the hypothalamic regions and inhibits the cortex, however, Fig 1 shows neither an activation of the hypothalamus in the ISO nor an inhibition of the cortex when compared to home cage control. If anything it suggests the opposite.

Thank you for your insightful observations regarding the discrepancies between Figures 2 and 3. We believe that when you refer to Figure 1, you are actually referencing Figure 2C.

ISO activation in Hypothalamus: In Figure 2C, we regret the oversight where we inadvertently interchanged the positions of ISO and Saline. When accurately represented, Figure 2C indeed shows that ISO notably activates the periventricular zone (PVZ) and the lateral zone (LZ) of the hypothalamus compared to the home cage group. Moreover, there's a discernible difference in the hypothalamic response between ISO and KET.

ISO's Effect on the Cortex: The main aim of Figure 3 was to highlight the differing responses between ISO and KET in the cortex. Notably, KET demonstrates a positive correlation with PC1 (+7 on PC1), whereas ISO shows a negative association (-3 on PC1). Given that the coefficient of PC1 for the cortical region is positive, it suggests that the cortical areas activated by KET are inhibited by ISO (with KET's distribution around 0 on PC2). However, the divergence between ISO and the home cage is most apparent in PC2, with ISO clusters at +4 and the home cage approximately at -2, suggesting that ISO activates a different set of cortical nuclei. In alignment with this, Figure 2C also illustrates that ISO activates specific cortical areas, such as ILA and PIR, in contrast to the home cage.

Thus, Figure 3 primarily employs PCA to delineate the contrasts between ISO and KET, whereas Figure 2C emphasizes the comparison of each against their respective controls.

1. Control for isoflurane should be air in the induction chamber rather than home cage. It is possible that Fos activation reflects handling/stress pre-anesthesia in the animals, which would increase Fos expression in the stress-related regions such as the BST, striatum (CeA), hypothalamus (PVH) and potentially the LC.

Thank you for emphasizing the importance of an appropriate control for Isoflurane.

In our efforts to minimize the potential impact of stress-induced c-Fos expression, we implemented several precautionary measures. Prior to the experiment, both groups of mice were subjected to handling and acclimatization within the induction chamber over four days. By the day of the experiment, for the mice in the experimental group, we ensured they were comfortable and exhibited no signs of distress or fear—such as cowering or evading. With care, we slowly relocated them to the nearby anesthesia induction chamber. Using 5% ISO, anesthesia was induced promptly, following a meticulously devised protocol to reduce stress impacts on c-Fos expression.

Moreover, existing studies have shown Isoflurane's activation of BST/CeA (Hua, T, et al., Nat Neurosci, 2020, 23: 854-868), PVH (Xu, Z, et al., British Journal of Anaesthesia, 2023, 130: 446-458), and LC (Lu, J, et al., J Comp Neurol, 2008, 508: 648-62), even when using oxygen controls. Such literature supports our findings, indicating that the activation we observed was indeed due to Isoflurane and not purely stress-related.

1. In the Ket network there are a few anticorrelated regions, most of which are amongst the list of the most activated regions, does this mean that the strong correlation results from an overall decreased activation? And if so, is it possible that the ketamine anesthesia was stronger than the isoflurane, causing a more general reduction in activity?

The pronounced correlations observed within the ketamine (KET) network do not signify a generalized decrease in activation. Instead, these correlations reflect significantly enhanced activity in specific regions under KET anesthesia. This amplified correlation is an indication of a more widespread increase in activity, rather than a decrease. These findings are consistent with previous research, which showed that anesthetic doses of ketamine produce patterns of Fos expression in the CNS similar to wakefulness (Lu, J, et al., J Comp Neurol, 2008; 508(4): 648-62).

Regarding the comparative strength of KET versus ISO anesthesia, our electroencephalographic evidence confirms that both agents induce a loss of consciousness. No significant differences were observed in EEG and EMG readings within the first 30 minutes post-administration. In future research, a continuous intravenous or intraperitoneal administration of KET might be a preferable method.

1. Since they have established networks it would be easy and useful to look at how the different regions identified (sleep, pain, neuroendocrine, motor-related, ...) work together to maintain analgesia, are they within the same module? Do they become functionally connected and is this core network of functional connections similar for KET and ISO?

Thank you for your suggestion. In response to your inquiry, we undertook analysis of the core functional networks for KET and ISO, using a set threshold at r>0.82 and P<0.05. For evaluating the modularity of each network, we utilized Newman's spectral community detection algorithm.

(A) The ISO’s core functional network (56 nodes, 372 edges) predominantly divides into two modules with a modularity quotient of 0.345. ISO-active regions include arousal-associated regions (PL, ILA, PVT), analgesia-related (CeA, LC, PB), neuroendocrine function nuclei (TU, PVi, ARH, PVH, SON) as detailed in Figure 5. Notably, ARH and SON weren't incorporated into the core network. Analgesia-associated regions, such as CeA, LC, and PB, reside within module 1, while neuroendocrine nuclei are spread between modules 1 and 2.

(B) In contrast, KET's core functional network (61 nodes, 1820 edges) splits into three distinct modules, but its low modularity quotient (0.06) indicates a lack of clear functional modularization, suggesting denser interconnections among brain regions. Furthermore, functionally-related regions such as arousal (PL, ILA, PVT, DR), analgesia-related (ACA, APN, PAG, LC), and neuroendocrine regulation (PVH, SON),etc., as seen in Figure 4, are distributed across different modules. This distribution may implies that functions like analgesia and neuroendocrine regulation are not governed by simple, linear processes, but arise from complex, overlapping pathways spanning various modules and functional zones.

In summary, the core functional networks of ISO and KET differ, with functionally-related regions spanning multiple modules, reflecting their diverse roles in varied physiological regulations.

1. The naming of the function of some of the regions is very much debatable. For instance, PL/ILA are named "sleep-wakefulness regulation" regions in the paper. I can think of many more important functions of the PL/IL including executive functions, behavioral flexibility, and emotional control. It is unclear how the functions of all the regions were attributed. I am not sure that this biased labeling of structure-function is useful to the reports, it may instead suggest wrong conclusions.

Thank you for your thoughtful feedback regarding our classification of the functions of the PL/ILA regions in our manuscript.

We recognize the challenge in accurately defining the functions of brain regions. While there is evidence highlighting the role of PL/ILA in arousal pathways, we also acknowledge their documented roles in executive functions, behavioral flexibility, and emotional control. In response to your comments, we have refined our description, changing "sleep-wakefulness regulation" to "wake-promoting pathways" (see Line: 159, 164).

It's worth noting that many brain regions, including the PL/ILA, have multiple functions. We agree that a single label might not capture the entirety of their roles. To provide a broader perspective, we will add a section in our manuscript that sheds light on the varied functions of these regions (Line: 181).

1. A point of concern and confusion is the number of brain regions analyzed. In the introduction, it is mentioned that 987 brain regions are considered, but this is reduced to 53 selected brain regions in Figure 2, then 201 brain regions in Figure 3, and reduced again to 63 for the network analysis. The rationale for selecting different brain regions is not clear.

For the 987 brain regions: Using the standard mouse atlas available at http://atlas.brain-map.org/, the mouse brain is organized into nine levels. The broadest category is the grey matter, which then progresses to more specific subdivisions, totaling 987 unique regions.

For the 53 brain regions: To effectively understand the activation patterns of ISO and KET, we started with a broad approach, looking at larger brain areas like the thalamus and hypothalamus. This broad view, presented in Figure 2, focuses on the 5th-level brain regions, encompassing 53 primary areas. This methodology is also employed in the study by Do et al. (Elife, 2016; 5: e13214). We have added the rationale for selecting these brain regions in the main text (Line: 92).

Regarding the 201 brain regions in Figures 3, 4, and 5: We delved deeper, examining the 6th-level brain regions, a common granularity in neuroscience research. This detailed view allowed us to highlight specific areas, like the CeA and PVH (Line:129).

Finally, for Figures 6 and 7, we selected 63 regions that were activated by both ISO and KET, as well as regions previously reported to be related to the mechanism of general anesthesia(Leung, L, et al., Progress in neurobiology, 2014; 122: 24-44) (Line: 220). Using these regions, we analyzed the correlation of c-Fos expression, aiming to construct a functional brain network with strong positive connections.

We hope this clarifies our approach and the rationale behind our region selection at each stage of the study. Thank you for your attention to this detail.

1. The statistical analysis does not seem appropriate considering the high number of comparisons. They use simple t-tests without correction for multiple comparisons.

Thank you for pointing out the concern regarding our statistical analysis. In the revised manuscript, we addressed the issue of multiple comparisons correction in our t-tests. We adopted the statistical methods detailed in the papers by Renier, N, et al., Cell, 2016; and Benjamini, Y, and Y Hochberg, 1995. P-values were adjusted for multiple comparisons using the two-stage linear step-up procedure of Benjamini, Krieger, and Yekutieli, with a false discovery rate (FDR) threshold (Q) of 0.05. This approach is now explained in the Materials and Methods section (Line: 434). After this adjustment, the brain regions we initially identified remained statistically significant. Furthermore, we revisited the original immunohistochemical images to confirm the differences in c-Fos cell expression between the experimental and control groups, reinforcing our conclusions.

1. There is no statistical analysis in Fig 2C。

Thank you for bringing to our attention the lack of statistical analysis in Fig 2C. We have now added the relevant statistical data in Supplementary Table 1 and provided annotations in Fig 2C to reflect this.

Reviewer #2

1. The authors report 987 brain regions in the introduction, but I cannot find any analysis that incorporates these or even which regions they are. Very little rationale is provided for the regions included in any of the analyses and numbers range from 53 in Figure 1, to 201 in Figure 3, to 63 in Figure 6. It would help if the authors could first survey Fos+ counts across all regions to identify a subset that is of interest (significantly changed by either condition compared to control) for follow up analysis.

Thank you for your insightful comments on the number of brain regions analyzed in our study.

987 Brain Regions: The reference to 987 brain regions from the standard mouse atlas (http://atlas.brain-map.org/) represents the entire categorization of the mouse brain across nine levels. We recognize that a comprehensive analysis of all these regions would be valuable, but to ensure clarity and depth, we took a focused approach.

Region Selection Rationale:

Figure 2: Concentrated on 5th-level brain regions (53 areas), inspired by methods from Do et al. (eLife, 2016;5:e13214). This provided a broad overview of c-Fos expression differences. Figures 4 and 5: Delved into 6th-level brain regions (201 areas), a common practice in neuroscience for more detailed study. Figure 6: We focused on 63 regions, which encompass not only the regions activated by both ISO and KET but also those previously reported to be associated with the mechanisms of general anesthesia. Methodological Approach: Our region selection was rooted in identifying areas with significant changes under anesthetic conditions compared to controls. This staged approach allowed a targeted analysis of the most affected regions, ensuring robust conclusions.

Enhancements: We've incorporated comparative analyses of activated brain regions at different hierarchical levels in Figures 4 and 5. For clearer comprehension, we’ve added clarifications in the manuscript at Lines: 92, 130, and 220.

1. Different data transformations are used for each analysis. One that is especially confusing is the 'normalization' of brain regions by % of total brain activation for each animal prior to PCA analysis in Figures 2 and 3. This would obscure any global differences in activation and make it unlikely to observe decreases in activation (which I think is likely here) that could be identified using the Fos+ counts after normalizing for region size (ie. Fos+ count / mm3) which is standard practice in such Fos-based activity mapping studies. While PCA can be powerful approach to identify global patterns, the purpose of the analysis in its current form is unclear. It would be more meaningful to show that regional activation patterns (measured as counts/mm3) are on separate PCs by group.

Thank you for your thoughtful comments. We regret any confusion caused by our initial presentation. For the PCA analysis in Figures 2A and 3A, we calculated the ratio of cell density in each brain region to the overall brain density, and then applied a logarithmic transformation to this ratio. Our approach in Figure 2C was to use the proportion of c-Fos cell counts in individual brain regions to the total cell counts throughout the brain. This methodology considers variations in overall c-Fos cell counts across animals, effectively mitigating potential biases due to differential global activation levels across subjects.

Furthermore, our direct comparison of differences in c-Fos cell counts between ISO, KET, and their respective control groups in Figures 4 and 5 addresses your concerns about potential decreases in activation. Notably, we did not identify any brain regions with significant suppression in these figures, which is consistent with the trends observed post-normalization in Figure 2C.

Given your feedback, we conducted another PCA using cell densities for each region (counts/mm3). However, we found significant variability and non-normal distribution of c-Fos density across the groups, leading to extensive data dispersion. Consequently, normalizing the cell counts across regions and then applying a logarithmic transformation before PCA might be more appropriate.

Additionally, our exploration of regional activation patterns using PCA analysis for ISO and KET separately, based on the logarithm ratio of the c-Fos density, revealed that there was no distinct clustering feature among the different brain regions (as illustrated in Author response image 5: colors represented distinct brain regions, while the shapes were indicative of different clusters). This observation further suggests that our original statistical approach might be more suitable.

1. Critical problem: The authors include a control group for each anesthetic (ketamine vs. saline, isofluorane vs. homecage) but most analyses do not make use of the control groups or directly compare Fos+ counts across the groups. Strictly speaking, they should have compared relative levels of induction by ketamine versus induction by isoflurane using ANOVAs. Instead, each type of induction was separate from the other. This does not account for increased variability in the ketamine versus isoflurane groups. There is no mention in the Statistics section or in Results section that any multiple comparison corrections were used. It appears that the authors only used Students t-test for each region and did not perform any corrections.

We appreciate the reviewer's insights and have addressed your concerns:

Given the pronounced difference in c-Fos cell count expression between the KET and ISO groups, a direct comparison of Fos+ counts may not effectively capture their inherent disparities. To better highlight these distinctions, we used the logarithm ratio of c-Fos density in our PCA analysis (Figure 3), mitigating potential disparities in overall cell counts between samples and emphasizing relative variations. However, in response to your feedback, we've included additional analyses. Author response image 6 depicts the c-Fos density (cells/mm^3) across different brain regions for the home cage, ISO, saline, and KET groups, with regions like the cerebral cortex, cerebral nuclei, thalamus, and others differentiated by shaded backgrounds. Data are represented as mean ± SEM. We performed a one-way ANOVA followed by Tukey’s post hoc test, marking significant differences between ISO and KET with asterisks: ***P < 0.001, **P < 0.01, *P < 0.05.

Regarding multiple comparison corrections, we've conducted thorough analyses on the data in Figure 2C and Figures 4, 5, and 6, implementing multiple comparison corrections. The detailed methodology is provided in the “Statistical analysis” section.

1. Figures 4 and 5 show brain regions 'significantly activated' following KET or ISO respectively, but again a subset of regions are shown and the stats seem to be t-tests with no multiple comparisons correction. It would help to show these two figures side by side, include the same regions, and keep the y axis ranges similar so the reader can easily compare the 'activation patterns' across the two treatments. Indeed, it looks like KET/Saline induced activation is an order or magnitude or two higher than ISO/Homecage. I would also recommend that this be the first data figure before any other analyses and maybe further analysis could be restricted to regions that are significantly changed in following KET or ISO here.

Thank you for your constructive feedback regarding Figures 4 and 5.

Comparison and Presentation of Figures 4 and 5: We acknowledge your suggestion to present these figures side by side for easier comparison. In the supplementary figure provided in the previous question, we've placed Figures 4 and 5 adjacent to each other, with consistent y-axis ranges, ensuring that readers can make direct comparisons between the activation patterns elicited by KET and ISO.

Statistical Concerns and Region Selection: As mentioned in our previous response, we have conducted multiple comparison corrections on the data presented in Figures 4 and 5. Detailed procedures are elaborated in the “Statistical analysis” section. We believe this approach addresses your concerns regarding the use of t-tests without corrections for multiple comparisons.

Difference in Activation Levels: We observed that the c-Fos activation due to KET is significantly higher than that from ISO. When presented side-by-side using the same scale, ISO activations appear less prominent, potentially mask subtle differences in the activation patterns of ISO, particularly if both KET and ISO showed changes in the same direction in certain brain regions but differed in magnitude. To address this, we used the proportion of c-Fos cell counts in Figure 2C, the logarithm ratio of c-Fos density in Figure 2A and Figure 3. This method emphasizes the relative changes, rather than absolute values, giving a more balanced view of the effects of each treatment.

1. Analyses in Figure 6 and 7 are interesting but again the choice of regions to include is unclear and makes interpreting the results impossible. For example, in Figure 7 it is unclear why the list of regions in bar graphs showing Degree and Betweenness Centrality are not the same even within a single row?

Thank you for your pertinent observation. The choice of brain regions in Figures 6 and 7 was carefully determined based on two main criteria: regions that were significantly activated by ISO or KET within the scope of our study, and those previously reported to be associated with anesthesia mechanisms and sleep-wake regulation.

Regarding your second concern on Figure 7, the discrepancies observed in the x-axes of the bar graphs arise from our methodological approach. We prioritized presenting the top 20% of regions based on their Degree or Betweenness Centrality values. By separately ranking these regions from highest to lowest, the regions presented for each metric inherently differ. This approach was taken to elucidate nodes that consistently emerge as significant across both metrics, thereby highlighting core nodes in the functional network. Were we to use a consistent x-axis without this ranking, it would not only necessitate a more extensive presentation but might also dilute the emphasis on key information. To clarify this methodology and its rationale for our readers, we have expanded upon this in the manuscript at Line 243.

We hope these clarifications address your concerns and facilitate a clearer understanding of our findings.

Reviewer #1 (Recommendations For The Authors):

Minor points

1. In Table 1: the separation of which substructures belong to which brain structure is not clear

1. Line 132 on page 3 seems to repeat the sentence earlier in the paragraph "KET predominantly affects brain regions within the cerebral cortex (CTX), while significantly inhibiting the hypothalamus, midbrain, and hindbrain."

1. Typos

a) Line 99/100 and 130 Central nucleus (CNU) should be cerebral nucleus

b) Comma on line 166

c) Fig. 4D: KET instead of Keta

d) Line 263 "ep"

e) Line 332: 35" "ml (add space)

1. Will data and code be made available?

Thank you for your detailed feedback.

1. We have revised Table 1 to clarify which substructures belong to which brain structures.

2. We acknowledge the redundancy and have now edited line 139 on page 3 to remove the repeated sentence regarding the effects of KET on brain regions.

3. We have addressed the typos you pointed out:

a. The terms "Central nucleus (CNU)" have been corrected to "cerebral nucleus."

b. The comma issue on line 166 has been rectified.

c. In Fig. 4D, we have corrected "Keta" to "KET."

d. We have corrected the typo "ep" on line 263.

e. A space has been added between "35" and "ml" on line 332 as you indicated.

1. Regarding the availability of data and code, we are currently conducting additional analyses related to this study. Once these analyses are completed, we will be more than happy to make the data and code available.

Thank you for assisting us in improving our manuscript.

Reviewer #2 (Recommendations For The Authors):

Minor comments:

1. The term 'whole-brain mapping' in the title suggests that the mapping was performed on 'intact brains' where in fact serial sections were used here. Maybe the authors could change to 'brain-wide mapping' to align better with the study.

Thank you for your insightful comments.

We have revised the title as suggested, changing "whole-brain mapping" to "brain-wide mapping".

1. It is unclear if the mice were kept under anesthesia for the 90-min duration and how the authors monitored the level of sedation. Additionally, if the KET mice were already sedated why were they further sedated with ISO before perfusions and tissue extraction? The methods should be clarified and any potential confounds discussed.

To maintain consistency in the experimental protocol and to reduce stress reactions in the mice, ISO was used before perfusion in all cases. However, this does not affect c-Fos expression as the expression of c-Fos protein starts 20-30 minutes after stimulation (Lara Aparicio, S Y, et al., NeuroSci, 2022; 3(4): 687-702).

We appreciate your guidance in enhancing the clarity of our manuscript.

Reviewer #3 (Recommendations For The Authors):

Recommendation: Minor corrections.

1. The authors should delve deeper into the molecular mechanisms underlying the observed effects, particularly the changes associated with NMDA and GABA receptors. Exploring these mechanisms would provide a more comprehensive understanding of how Ketamine and Isoflurane modulate neural activity and induce anesthesia.

1. The clinical relevance of these findings has not been sufficiently addressed. It would be valuable to elaborate on how the current research outcomes could potentially lead to changes in current anesthesia practices. For instance, identifying the distinct pathways of action for Ketamine and Isoflurane could aid anesthesiologists in selecting the most appropriate anesthetic based on the specific needs of individual patients or surgical procedures.

1. Both Ketamine and Isoflurane have been associated with neurotoxicity. It is important to discuss how the c-Fos activation induced by these anesthetics could contribute, at least partially, to anesthesia-related neurotoxicity. Examining the potential neurotoxic effects would provide a more comprehensive understanding of the risks associated with these anesthetics and aid in the development of safer anesthesia protocols.

Thank you for your valuable suggestions.

Regarding the three points (1, 2, and 3) you've raised, we fully recognize their significance. In the current study, our primary focus was on the differential impacts of Isoflurane and Ketamine on widespread c-Fos expression in the brain. However, we indeed acknowledge the importance of delving deeper into these mechanisms and their clinical relevance. Therefore, we intend to explore these critical issues in greater detail in our future research endeavors.

We appreciate your feedback, which provides constructive guidance for our subsequent research directions.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Ketamine (KET) and isoflurane (ISO) are two widely used general anesthetics, yet their distinct and shared neurophysiological mechanisms remain elusive. In this study, we conducted a comparative analysis of KET and ISO effects on c-Fos expression across the brain, utilizing principal component analysis (PCA) and c-Fos-based functional network analysis to evaluate the responses of individual brain regions to each anesthetic. Our findings demonstrate that KET significantly activates cortical and subcortical arousal-promoting nuclei, with the somatosensory cortex (SS) serving as a hub node, corroborating the top-down general anesthesia theory for dissociative anesthesia. In contrast, ISO activates the nuclei in the hypothalamus and brain-stem, with the locus coeruleus (LC) as a hub node, implying a bottom-up mechanism for anesthetic-induced unconsciousness. Notably, the coactivation of sleep-wakefulness regulation, analgesia-related, neuroendocrine-related nuclei (e.g., prelimbic area (PL) and infralimbic areas (ILA), and the anterior paraventricular nucleus (aPVT), Edinger-Westphal nucleus (EW), locus coeruleus (LC), parabrachial nucleus (PB), solitary tract nucleus (NTS)) by both anesthetics underscores shared features such as unconsciousness, analgesia, and autonomic regulation, irrespective of their specific molecular targets. In conclusion, our results emphasize the distinct actions of KET and ISO while also uncovering the commonly activated brain regions, thus contributing to the advancement of our understanding of the mechanisms under-lying general anesthesia.

Revised version: Ketamine (KET) and isoflurane (ISO) are two widely used general anesthetics, yet their distinct and shared neurophysiological mechanisms remain elusive. In this study, we conducted a comparative analysis of KET and ISO effects on c-Fos expression across the brain, utilizing principal component analysis (PCA) and c-Fos-based functional network analysis to evaluate the responses of individual brain regions to each anesthetic. Our findings demonstrate that KET significantly activates cortical and subcortical arousal-promoting nuclei, with the somatosensory cortex (SS) serving as a hub node, corroborating the top-down general anesthesia theory for dissociative anesthesia. In contrast, ISO activates the nuclei in the hypothalamus and brainstem, with the locus coeruleus (LC) as a hub node, implying a bottom-up mechanism for anesthetic-induced unconsciousness. Notably, the coactivation of sleep-wakefulness regulation, analgesia-related, neuroendocrine-related nuclei (e.g., prelimbic area (PL) and infralimbic areas (ILA), and the anterior paraventricular nucleus (aPVT), Edinger-Westphal nucleus (EW), locus coeruleus (LC), parabrachial nucleus (PB), solitary tract nucleus (NTS)) by both anesthetics underscores shared features such as unconsciousness, analgesia, and autonomic regulation, irrespective of their specific molecular targets. In conclusion, our results emphasize the distinct actions of KET and ISO while also uncovering the commonly activated brain regions, thus contributing to the advancement of our understanding of the mechanisms underlying general anesthesia.

Introduction

Earlier version: Despite considerable investigation into the molecular targets, neural circuits, and functional connectivity associated with various anesthetics, our comprehension of their effects on overall brain activity continues to be limited and incomplete [1]. At the molecular level, ketamine (KET) and isoflurane (ISO) interact with N-methyl-D-aspartate (NMDA) and gamma-aminobutyric acid type A (GABAa) receptors, respectively, modulating neuronal excitability and ultimately leading to a loss of consciousness [2]. In systems neuroscience, the neural mechanisms of anesthetic induced unconsciousness involve both top-down and bottom-up processes [3, 4]. As evidenced by in vivo electro-physiology or functional magnetic resonance imaging (fMRI) studies, the top-down paradigm illustrates that anesthetics induce unconsciousness by disrupting corticocortical and corticothalamic circuits responsible for neural information integration, while peripheral sensory information can still be conveyed to the primary sensory cortex [5, 6]. The bottom-up approach, exemplified by ISO, involves the activation of sleep-promoting nuclei like ventral lateral preoptic nucleus (VLPO) and inhibition of arousal centers in the brain-stem and diencephalon, supporting the shared circuits of sleep and anesthesia [7, 8]. However, the limited spatial resolution of fMRI studies and the inability of EEG to capture specific brainstem nuclei hinder the acquisition of comprehensive whole-brain information. Although a substantial body of knowledge has been amassed, our understanding of the reciprocal responses among different brain regions during general anesthesia remains relatively sparse and fragmented. To bridge these gaps, further investigation using advanced techniques that can capture the whole-brain dynamics is needed to elucidate the complex interactions and shared mechanisms between various anesthetics.

Revised version: Despite considerable investigation into the molecular targets, neural circuits, and functional connectivity associated with various anesthetics, our comprehension of their effects on overall brain activity continues to be limited and incomplete [1]. At the molecular level, ketamine (KET) and isoflurane (ISO) interact with N-methyl-D-aspartate (NMDA) and gamma-aminobutyric acid type A (GABAa) receptors, respectively, modulating neuronal excitability and ultimately leading to a loss of consciousness [2]. In systems neuroscience, the neural mechanisms of anesthetic induced unconsciousness involve both top-down and bottom-up processes [3, 4]. As evidenced by in vivo electrophysiology or functional magnetic resonance imaging (fMRI) studies, the top-down paradigm illustrates that anesthetics induce unconsciousness by disrupting corticocortical and corticothalamic circuits responsible for neural information integration, while peripheral sensory information can still be conveyed to the primary sensory cortex [5, 6]. The bottom-up approach, exemplified by ISO, involves the activation of sleep-promoting nuclei like ventral lateral preoptic nucleus (VLPO) and inhibition of arousal centers in the brainstem and diencephalon, supporting the shared circuits of sleep and anesthesia [7, 8]. However, the limited spatial resolution of fMRI studies and the inability of EEG to capture specific brainstem nuclei hinder the acquisition of comprehensive whole-brain information. Although a substantial body of knowledge has been amassed, our understanding of the reciprocal responses among different brain regions during general anesthesia remains relatively sparse and fragmented. To bridge these gaps, further investigation using advanced techniques that can capture the whole-brain dynamics is needed to elucidate the complex interactions and shared mechanisms between various anesthetics.

Earlier version: Neuronal extracellular stimulation typically results in the elevation of adenosine 3’,5’-cyclic monophosphate (cAMP) levels and calcium influx, ultimately leading to the upregulation of immediate early genes (IEGs) such as c-fos [9, 10]. The translation product of c-fos, c-Fos protein, offers single-cell spatial resolution and has been utilized as a biomarker to identify anesthetic-activated brain regions [11]. Previous investigations of c-Fos expression throughout the brain demonstrated that GABAergic agents inhibited cortical activity while concurrently activating subcortical brain regions, including the VLPO, median preoptic nucleus (MnPO), lateral septal nucleus (LS), Edinger-Westphal nucleus (EW), and locus coeruleus (LC) [12–15]. In contrast, KET was shown to provoke wake-like c-Fos expression and intense augmentation of c-Fos expression in various brain regions at clinical dosages (75-100 mg/kg) [13]. It is important to note that these experiments administered KET at lights-on and GABAa receptor agonists at lights-off, potentially introducing circadian influences for direct comparison of ISO and KET. Moreover, it has been re-vealed that state of general anesthesia is not determined by activity in individual brain areas, but emerges as a global change within the brain. This change involves the activation of lateral habenular nucleus (LHb), VLPO, supraoptic nucleus (SON), and central amygdaloid nucleus (CeA), which are essential for anesthetics induced sedation, unconsciousness, or analgesia [7, 16–18]. However, brain-wide mapping and comparison of distinct anesthetic (KET and ISO) activated nuclei at a cellular level have not been fully elucidated.

Revised version: Neuronal extracellular stimulation typically results in the elevation of adenosine 3’,5’-cyclic monophosphate (cAMP) levels and calcium influx, ultimately leading to the upregulation of immediate early genes (IEGs) such as c-fos [9, 10]. The translation product of c-fos, c-Fos protein, offers single-cell spatial resolution and has been utilized as a biomarker to identify anesthetic-activated brain regions [11]. Previous investigations of c-Fos expression throughout the brain demonstrated that GABAergic agents inhibited cortical activity while concurrently activating subcortical brain regions, including the VLPO, median preoptic nucleus (MnPO), lateral septal nucleus (LS), Edinger-Westphal nucleus (EW), and locus coeruleus (LC) [12–15]. In contrast, KET was shown to provoke wake-like c-Fos expression and intense augmentation of c-Fos expression in various brain regions at clinical dosages (75-100 mg/kg) [13]. It is important to note that these experiments administered KET at lights-on and GABAa receptor agonists at lights-off, potentially introducing circadian influences for direct comparison of ISO and KET. Moreover, it has been revealed that state of general anesthesia is not determined by activity in individual brain areas, but emerges as a global change within the brain. This change involves the activation of lateral habenular nucleus (LHb), VLPO, supraoptic nucleus (SON), and central amygdaloid nucleus (CeA), which are essential for anesthetics induced sedation, unconsciousness, or analgesia [7, 16–18]. However, brain-wide mapping and comparison of distinct anesthetic (KET and ISO) activated nuclei at a cellular level have not been fully elucidated.

In this study, we examined the distribution of nuclei activated by ISO and KET in 987 brain regions using immunochemical labeling and a customized MATLAB software package, which facilitated signal detection and registration to the Allen Mouse Brain Atlas reference [19]. We compared whole-brain c-Fos expression induced by KET and ISO through principal component analysis (PCA) and calculated inter-regional correlations by determining the covariance across subjects for each brain region pair. We then extracted significantly positively correlated brain regions to construct functional networks and performed graph theory-based network analyses to identify hub nodes. Our results uncovered distinct yet overlapping whole-brain activation patterns for KET and ISO.

Results

A comparison of the activation patterns of c-Fos in 53 brain areas in response to ISO and KET

Earlier version: To examine the pattern of c-Fos expression throughout the brain, 1.5% ISO was continuously ventilated, or 100 mg/kg KET was administrated 90 minutes before harvesting (Figure 1A). Raw images of brain slices were aligned to the Allen Mouse Brain Atlas (Figure 1B). The entire brain was divided into 12 main regions in grade four with a total of 53 subregions in supplementary table 1. We first calculated the c-Fos cell density in each brain region across the four groups and then log-transformed the data for principal component analysis (PCA). The top two principal components revealed a clear separation between the KET and ISO groups, with KET induced c-Fos expression closely associated with PC1, accounting for 30.72% of the variance, while ISO was associated with PC2, accounting for 25.89%. The absence of overlap between PC1 and PC2 suggests distinct patterns of brain region activation for each anesthetic (Figure 2A). Figure 2B illustrates the component coefficients of principal component 1 (PC1) and principal component 2 (PC2) for each brain region. The points in the figure represent brain regions with absolute component coefficients within the top 25%. KET, closely related to PC1, predominantly affected cortical regions, as evidenced by top 25% positive PC1 coefficients in MOB, AON, MO, ACA, and ORB, while subcortical areas exhibited negative coefficients. In contrast, ISO, showing a notable association with PC2, displayed top 25% positive coefficients in PALc, LSX of CNU, cortical ILA, and hypothalamic PVZ, with top 25% negative coefficients primarily found in cortical regions. These findings reveal distinct patterns of brain region activation for each anesthetic, with KET primarily affecting cortical regions, and ISO mainly influencing the central nucleus (CNU) and hypothalamus. To reduce inter-individual variability, we normalized c-Fos+ cell counts in each brain area by dividing them by the total c-Fos+ cells in the entire brain (Figure 2C), enabling a more accurate comparison of KET and SO induced activation patterns. In line with PCA results, KET significantly activated the isocortex (16.55 ± 2.31 vs. 40.13 ± 1.97, P < 0.001) and reduced the proportion of c-Fos+ cells in the hypothalamus (14.88 ± 2.18 vs. 4.03 ± 0.57, P = 0.001) and midbrain (21.87 ± 2.74 vs. 6.48 ± 1.06, P = 0.001) compared to the saline group. In contrast, ISO induced c-Fos+ cells were mainly observed in the striatum and hypothalamus compared to the home cage group, specifically in the periventricular zone (PVZ) (2.330 ± 0.759 vs. 6.63 ± 0.84, P = 0.004) and lateral zone (LZ) (2.82 ± 0.3 vs. 5.59 ± 0.75, P = 0.007) of the hypothalamus, as well as STRv (0.59 ± 0.10 vs. 2.77 ± 0.40, P < 0.001) and LSX (1.93 ± 0.59 vs. 5.11 ± 0.41, P = 0.0012) of the striatum (Figure 2C and Table 1). In summary, our findings reveal that KET predominantly affects cortical regions, whereas isoflurane mainly targets subcortical areas, specifically the striatum, and hypothalamus. These results demonstrate distinct patterns of brain region activation for each anesthetic agent.

Revised version: To examine the pattern of c-Fos expression throughout the brain, 1.5% ISO was continuously ventilated, or 100 mg/kg KET was administered 90 minutes before harvesting (Figure 1A). Concurrently, to quantitatively assess brain states across the four groups, we utilized electroencephalography (EEG) and electromyography (EMG) recordings (Figure 1—figure supplement 1). The depth of anesthesia was evaluated based on the power ratio of the EEG within the delta (0.5–4 Hz) and theta (6–10 Hz) frequency bands, combined with the changes in the EMG [20]. Our findings showed that ISO and KET notably elevated EEG power ratio while reducing EMG power, suggesting both drugs induced a loss of consciousness. The raw images of brain slices were aligned to the Allen Mouse Brain Atlas (Figure 1B). Based on the standard mouse atlas available at http://atlas.brain-map.org/, the mouse brain was segmented into nine hierarchical levels, totaling 987 regions. The primary level consists of grey matter, the secondary of the cerebrum, brainstem, and cerebellum, and the tertiary includes regions like the cerebral cortex and cerebellar nuclei, among others. Some regions extend to the 8th and 9th levels. Our initial statistical comparison was conducted on the fifth level, which is composed of 53 subregions, as detailed in Supplementary Table 1 (Figure 1B) [21]. To differentiate c-Fos expression patterns across groups, we calculated the c-Fos density for each brain region, divided it by the overall brain c-Fos density, and then applied a logarithmic transformation. The transformed values were used for principal component analysis (PCA) (Figure 2A). The KET and ISO groups were distinctly delineated by the top two principal components. Notably, the KET-induced c-Fos expression manifested a positive association with principal component 1 (PC1), clustering around +2.5, which represented 30.72% of the total variance. In contrast, it exhibited a negative correlation with principal component 2 (PC2), clustering near −2.0, which contributed to 25.89% of the variance. On the other hand, ISO displayed a positive correlation with PC2, clustering at approximately +2.5, while its association with PC1 was neutral, near 0. Figure 2B illustrates the component coefficients of PC1 and PC2 for each brain region. Dots represent the top 25% of brain regions based on the absolute values of their component coefficients, signifying the most influential regions. KET, predominantly positively correlated with PC1, is associated with the activation of cortical regions (as evidenced by positive PC1 coefficients in MOB, AON, MO, ACA, and ORB) and the inhibition of subcortical areas (indicated by negative coefficients). On the other hand, ISO, with its notable association with PC2, showed top 25% positive coefficients in PALc, LSX of CNU, cortical ILA, and hypothalamic PVZ, while its negative coefficients were found primarily in cortical regions. These observations underscore that KET and ISO exhibit distinct patterns of brain region activation. KET displays cortical activation and subcortical inhibition, whereas ISO shows a contrasting preference, activating the cerebral nucleus (CNU) and the hypothalamus while inhibiting cortical areas. To reduce inter-individual variability, we normalized c-Fos+ cell counts in each brain area by dividing them by the total c-Fos+ cells in the entire brain (Figure 2C), Subsequently, we directly compared the differences between ISO and KET with their respective control groups. In line with PCA results, KET significantly activated the isocortex (16.55 ± 2.31 vs. 40.13 ± 1.97, P < 0.001) and reduced the proportion of c-Fos+ cells in the hypothalamus (14.88 ± 2.18 vs. 4.03 ± 0.57, P = 0.001) and midbrain (21.87 ± 2.74 vs. 6.48 ± 1.06, P = 0.005) compared to the saline group. In contrast, the proportion of c-Fos+ cells induced by ISO were mainly observed in the striatum and hypothalamus compared to the home cage group, specifically in the periventricular zone (PVZ) (2.330 ± 0.759 vs. 6.63 ± 0.84, P = 0.015) and lateral zone (LZ) (2.82 ± 0.3 vs. 5.59 ± 0.75, P = 0.002) of the hypothalamus, as well as STRv (0.59 ± 0.10 vs. 2.77 ± 0.40, P < 0.002) and LSX (1.93 ± 0.59 vs. 5.11 ± 0.41, P = 0.004) of the striatum (Figure 2C and Table 1). In summary, our findings reveal that KET predominantly affects cortical regions, whereas isoflurane mainly targets subcortical areas, specifically the striatum, and hypothalamus. These results demonstrate distinct patterns of brain region activation for each anesthetic agent.

Brain-wide quantification of c-Fos expression.

(A) Schematic representation of the habituation protocol typically used to acclimate mice. After being exposed to anesthetics for 90 minutes, the mice were euthanized. (B) Steps for data processing. Example of brain section registration to a corresponding coronal section from the Allen Brain Atlas. For Atlas rotation, the Allen reference atlas was rotated to mimic the slice angle of the experimental brain. Image registration maps the original coronal image (upper panel) to the corresponding Allen mouse brain atlas (lower panel). The registration module applies several geometric transformations (translation, rotation, and scaling) to optimize the matching of the original image to the anatomical structures. Fluorescence signals were detected from the original image (upper panel), and once detected, they were projected onto the Allen Mouse Brain Atlas for quantification and network analysis by means of the detected signals labeled with yellow boxes.

Whole-brain distributions of c-Fos+ cells induced by ISO, KET, and control conditions.

Earlier version: (A) Scatter plot of four conditions in the space spanned by the first two principal components (PC1 versus PC2) of c-Fos density from 53 brain regions. (B) Line plot of coefficients of PC1 and PC2. Dots represent regions with absolute values larger than 75 percentiles. (C) The normalized c-Fos+ cells in 53 brain areas (Home cage, n = 6; ISO, n= 6 mice; Saline, n = 8; KET, n = 6). Brain areas are grouped into 12 generalized, color-coded brain structures. Abbreviations of the 53 brain areas and percentages of c-Fos+ cells are listed in supplementary table 1. Error bar, mean ± SEM.

Revised version: (A) Scatter plot of four conditions in the space spanned by the first two principal components (PC1 versus PC2) of c-Fos density from 53 brain regions. (B) Line plot of coefficients of PC1 and PC2. Dots represent regions with absolute values larger than 75 percentiles. (C) The normalized c-Fos+ cells in 53 brain areas (Home cage, n = 6; ISO, n= 6 mice; Saline, n = 8; KET, n = 6). Brain areas are grouped into 12 generalized, color-coded brain structures. Abbreviations of the 53 brain areas and percentages of c-Fos+ cells are listed in supplementary table 1. Error bar, mean ± SEM. Mann-Whitney tests were corrected using the false discovery rate approach with a Q-value set at 0.05. Home cage vs ISO: *P<0.05, **P<0.01, ***P<0.001. Saline vs KET: #, P<0.05, ##, P<0.01, ###, P<0.001.

Similarities and differences in ISO and KET activated c-Fos brain areas

Earlier version: To further elucidate the similarities and differences between ISO and KET induced c-Fos expression patterns, we conducted a PCA on the logarithm of cell density across 201 brain regions, aiming to analyze their responses and characterize the variance among the respective experimental groups (Figure 3). The top two principal components displayed well-separated ISO and KET groups (Figure 3A), where the KET group had a strong association with PC1 (accounting for 28.97% of the variance) and the ISO group was associated with PC2 (accounting for 15.54% of the variance), indicating distinct features for the ISO and KET groups. We then further calculated the PC coefficients corresponding to the c-Fos density within each brain region (Figure 3B). The points in the figure represent brain regions with absolute component coefficients within the top 25%, highlighting the most influential regions in the PCA analysis. KET, closely associated with PC1, mainly activates brain regions within the cerebral cortex (CTX), including the visceral area (VISC), claustrum (CLA), dorsal peduncular area (DP), orbital area (ORB), and temporal association areas (TEA), while significantly inhibiting the hypothalamus, midbrain, and hindbrain. In contrast, ISO, closely associated with PC2, predominantly activates the hypothalamic regions, such as the ventrolateral preoptic nucleus (VLPO), suprachiasmatic nucleus (SCH), tuberal nucleus (TU), medial preoptic area (MPO), and supraoptic nucleus (SON), as well as certain nuclei in the central nucleus (CNU), while suppressing the cortex, midbrain, and hindbrain. KET predominantly affects brain regions within the cerebral cortex (CTX), while significantly inhibiting the hypothalamus, midbrain, and hindbrain. This result is in line with the top-down mechanism proposed previously, whereby anesthetic agents suppress consciousness by modulating cortical and thalamocortical circuits involved in the integration of neural information. ISO primarily activates hypothalamic regions, while suppressing the cortex, midbrain, and hindbrain. This aligns with the bottom-up mechanism, where anesthetics suppress consciousness by modulating sleep-wakefulness nuclei and neural circuits in the brainstem and diencephalon that have evolved to control arousal states[4]. We summarized the responses differences between the two anesthetics in Figure 3C, KET acts mainly via the top down mechanism, affecting cortical regions, while ISO operates through the bottom up mechanism, primarily targeting subcortical areas such as the hypothalamus, suggesting that different anesthetic agents may achieve the loss of consciousness by modulating distinct brain regions and neural circuits.

Revised version: To elucidate the c-Fos expression patterns induced by ISO and KET more clearly, we refined our focus from the initial 53 brain regions to a more granular examination of subregions, culminating in a comprehensive assessment of 201 distinct areas. We conducted a PCA on the logarithmic ratio of c-Fos density across these regions to discern variance among the groups (Figure 3). The KET group predominantly clustered around +7 on PC1 and approximately 0 on PC2, whereas the ISO group clustered around +4 on PC2 and −3 on PC1. KET displayed a strong positive correlation with PC1, accounting for 28.97% of the variance. In contrast, ISO negatively correlated with PC1 and positively with PC2, which accounted for 15.54% of the variance (Figure 3A). These findings underscore the distinct features between the ISO and KET groups, with the top two principal components effectively separating them. Further analysis of PC coefficients in each brain region showed (Figure 3B) that KET, positively associated with PC1, mainly activates the cerebral cortex (CTX) areas including the visceral area (VISC), claustrum (CLA), dorsal peduncular area (DP), orbital area (ORB), and temporal association areas (TEA), while inhibiting the hypothalamus, midbrain, and hindbrain. In contrast, ISO is negatively correlated with PC1, suggesting that the cortical areas activated by KET are inhibited by ISO. Additionally, ISO has a positive association with PC2 and is predominantly related to the activation of hypothalamic regions, such as the ventrolateral preoptic nucleus (VLPO), suprachiasmatic nucleus (SCH), tuberal nucleus (TU), medial preoptic area (MPO), and supraoptic nucleus (SON), as well as certain nuclei in the central nucleus (CNU). KET primarily activates regions within the cerebral cortex (CTX) and inhibits areas within the hypothalamus, midbrain, and hindbrain, aligning with the top-down mechanism suggesting anesthetic agents influence cortical and thalamocortical circuits [4]. Conversely, ISO activates hypothalamic areas while suppressing the cortex, midbrain, and hindbrain, consistent with the bottom-up mechanism. Figure 3C summarizes these distinctive responses, implying that anesthetic agents like KET and ISO modulate consciousness by targeting particular brain regions and neural circuits.

Similarities and differences in ISO and KET activated c-Fos brain areas.

(A) Scatter plot of four conditions in the space spanned by the first two principal components (PC1 versus PC2) of normalized c-Fos density from 201 brain regions. (B) Line plot of coefficients of PC1 and PC2. Dots represent regions with absolute values larger than 75 percentiles. (C) Activation patterns of ISO and KET in different parts of brain regions were summarized. ISO was predominantly characterized by increased activation of hypothalamic and cerebral nuclei while reducing activity in the cortex and midbrain, indicating that ISO primarily exerts its effects through a bottom-up mechanism. On the other hand, KET primarily activated the cortex and inhibited hypothalamus and midbrain activity, suggesting that KET regulates consciousness through a top-down mechanism.

Earlier version: Identifying brain regions activated by KET

Revised version: Identification of brain regions activated by KET

Earlier version: We further employed quantitative methods to identify the brain regions activated by KET across the entire brain. Our findings concur with previous studies, numerous cortical regions associated with somatosensory, auditory, visual, and movement were activated (Figure 4A and C, Supplementary Figure 1, Supplementary Table 2). Additionally, we identified several innovative observations that enrich the current understanding in this field. To provide a clearer overview of our findings, we categorized the activated brain areas based on their functions: (1) Arousal: Several nuclei associated with arousal were activated,includingtheprelimbicarea(PL),infralimbiccortex(ILA), paraventricular nucleus of the thalamus (PVT), sublaterodorsal nucleus (SLD), and dorsal raphe (DR). Prior evidence indicates that the PL/ILA and PVT regions play a role in regulating sleep and arousal in both cortical and subcortical areas [20]. The SLD has been reported to stabilize REM sleep, while dopaminergic neurons in the DR are involved in sleep-wake regulation [21, 22]. (2) Pain modulation, KET significantly activated pain-related areas such as the anterior cingulate cortex (ACA) in the cortex [23], the anterior pretectal nucleus (APN) in the thalamus, which is known to be involved in managing chronic pain [24], and the anterior periaqueductal gray (PAG) region and medial prefrontal cortex (mPFC), both part of the endogenous pain inhibitory pathway [25, 26],. Additionally, the activation of the locus coeruleus (LC) in the midbrain may contribute to KET’s analgesic effects [27]. (3) Neuroendocrine regulation: The paraventricular hypothalamic nucleus (PVH) and supraoptic nucleus (SON), which are neuroendocrine-related regions, were also activated [18, 28]. (4) Movement: Subcortical nuclei associated with movement, such as the subthalamic nucleus (STN) and nucleus incertus (NI), were prominently activated by KET administration [29, 30]. (5) Connectivity: We observed significant activation of the nucleus reunions (Re) located in the thalamus, which receives substantial afferent input from limbic structures and serves as a connector linking the hippocampus to the medial prefrontal cortex [31]. In summary, our study identified extensive activation of cortical and subcortical nuclei during KET anesthesia, encompassing regions related to arousal, pain modulation, neuroendocrine regulation, movement, and connectivity.

Revised version: We further employed quantitative methods to identify the brain regions activated by KET across the entire brain. Our findings concur with previous studies: numerous cortical regions associated with somatosensory, auditory, visual, and movement were activated (Figure 4A and C, Supplementary Figure 1, Supplementary Table 2). Additionally, we identified several innovative observations that enrich the current understanding in this field. To provide a clearer overview of our findings, we categorized the activated brain areas into different functional groups: (1) Arousal and REM sleep regulation: Several nuclei associated with arousal were activated, including the prelimbic area (PL), infralimbic cortex (ILA), paraventricular nucleus of the thalamus (PVT), dorsal raphe (DR), and sublaterodorsal nucleus (SLD). Prior evidence indicates that the PL/ILA and PVT regions play a role in regulating arousal in both cortical and subcortical areas[22, 23]. Additionally, dopaminergic neurons within the DR have been identified as vital components of wake-promoting pathways [24]. In the context of REM sleep, the sublaterodorsal nucleus (SLD) stands out for its role in stabilization[25]. (2) Pain modulation, KET significantly activated pain-related areas such as the anterior cingulate cortex (ACA) in the cortex [26], the anterior pretectal nucleus (APN) in the thalamus, which is known to be involved in managing chronic pain [27], and the anterior periaqueductal gray (PAG) region and medial prefrontal cortex (mPFC), both part of the endogenous pain inhibitory pathway [28, 29]. Additionally, the activation of the locus coeruleus (LC) in the midbrain may contribute to KET’s analgesic effects [30]. (3) Neuroendocrine regulation: The paraventricular hypothalamic nucleus (PVH) and supraoptic nucleus (SON), which are neuroendocrine-related regions, were also activated [18, 31]. (4) Movement: Subcortical nuclei associated with movement, such as the subthalamic nucleus (STN) and nucleus incertus (NI), were prominently activated by KET administration [32, 33]. (5) Connectivity: We observed significant activation of the nucleus reuniens (Re) located in the thalamus, which receives substantial afferent input from limbic structures and serves as a connector linking the hippocampus to the medial prefrontal cortex [34]. While our functional categorization offers a simplified overview, it’s worth noting that each activated brain region may have multifaceted roles beyond this classification. In summary, our study identified extensive activation of cortical and subcortical nuclei during KET anesthesia, encompassing regions related to arousal, pain modulation, neuroendocrine regulation, movement, and connectivity.

c-Fos expression in distinct brain regions after exposure to KET.

Earlier version: (A) Representative immunohistochemical staining of MOB, AON, ORB, MPO, ACA, MO, TRS, PL, ILA, DP, LS, PVT, SO, PVH, RE, VISC, AI, CLA, EPd, PIR, COA, AUD, TEa, ECT, PERI, CeA, SS, DG, STN, RSP, APN, LAT, EW, DR, PAG, SLD, PB, TRN, NI, LC, NTS, and NI c-Fos+ cells from the indicated mice. Scale bar, 200 µm. (B) Cell counts of saline group and KET group compared with P values < 0.05. Data are represented as mean ± SEM. (C) Schematic cross-section of the mouse brain showing activated brain regions by KET. Different colors indicate distinct functional nuclei. The red nuclei are associated with the regulation of sleep-wakefulness, the blue-green nuclei are linked to analgesia, the yellow nuclei are associated with motor function, and the white nuclei are a composite of various functional nuclei.

Revised version: (A) Representative immunohistochemical staining of MOB, AON, ORB, MPO, ACA, MO, TRS, PL, ILA, DP, LS, PVT, SO, PVH, RE, VISC, AI, CLA, EPd, PIR, COA, AUD, TEa, ECT, PERI, CeA, SS, DG, STN, RSP, APN, LAT, EW, DR, PAG, SLD, PB, TRN, NI, LC, NTS, and NI c-Fos+ cells from the indicated mice. Scale bar, 200 µm. (B) Representation of cell counts in brain regions exhibiting significant statistical differences. Data are shown as mean ± SEM. P values < 0.05. (C) Schematic cross-section of the mouse brain showing activated brain regions by KET. Different colors indicate distinct functional nuclei. The red nuclei are associated with the regulation of sleep-wakefulness, the blue-green nuclei are linked to analgesia, the yellow nuclei are associated with motor function, and the white nuclei are a composite of various functional nuclei.

Earlier version: Identifying brain regions activated by ISO

Revised version: Identification of brain regions activated by ISO

Earlier version: In our study, we aimed to identify brain regions activated by ISO and compare these with those activated by KET. We began by summarizing previously reported ISO activated nuclei, including PIR, LSd/LSv, CeA in the cortex and striatum, and VLPO, MnPO, EW, NTS, LC, ventral group of the dorsal thalamus (VeN), and area postrema (AP) in the hypothalamus and midbrain (Supplementary Table 3). We subsequently conducted a comprehensive MATLAB-based analysis of c-Fos expression throughout the entire brain, uncovering previously undetected activated nuclei (Figure 5A, Supplementary Figure 2). Newly identified activated nuclei in the CTX and the CNU included PL/ILA and ENT, aPVT in the thalamus, TU, ARH, PVi, and PVH in the hypothalamus, and PB in the hindbrain. All nuclei activated by ISO in this study were functionally classified and depicted in Figure 5C. Our results confirmed the activation of several brain regions involved in sleep-wakefulness regulation, such as the prelimbic area (PL) and infralimbic areas (ILA), and the anterior paraventricular nucleus (aPVT), leads to increased NREM sleep and decreased wakefulness [20]. Additionally, we observed activation in previously reported analgesia-related nuclei, including CeA and LC, as well as the parabrachial nucleus (PB) [18, 27, 32]. We also found activation in neuroendocrine function-related nuclei of the hypothalamus, such as TU, PVi, ARH, PVH, and SON. Moreover, we identified activations related to ISO induced side effects, such as in the piriform cortex (PIR) [33] and ENT [34]. which may be stimulated by ISO odor, and the solitary tract nucleus (NTS), potentially responsible for ISO induced vomiting [35]. The only activated nucleus in the midbrain was the Edinger-Westphal nucleus (EW). Recent research has found that sevoflurane activates EW and is involved in sleep induction and maintenance of anesthesia, suggesting its crucial role in general anesthesia [36]. By comparing the ISO and KET induced c-Fos expression, we summarized the brain regions activated by both anesthetics in Figure 5D. Despite variations in molecular targets, the coactivation of regions such as PL/ILA, aPVT, CeA, PVH, SON, EW, PB, LC, and NTS by both ISO and KET suggests an overlapping neuronal circuitry that influences sleep-wake regulation, analgesia, and neuroendocrine functions. This shared neural circuitry may potentially offer a common mechanism across the two anesthetics for the maintenance of general anesthesia.

Revised version: In our study, we aimed to identify brain regions activated by ISO and compare these with those activated by KET. We began by summarizing previously reported ISO activated nuclei, including PIR, LSd/LSv, CeA in the cortex and striatum, and VLPO, MnPO, EW, NTS, LC, ventral group of the dorsal thalamus (VeN), and area postrema (AP) in the hypothalamus and midbrain (Supplementary Table 3). We subsequently conducted a comprehensive MATLAB-based analysis of c-Fos expression throughout the entire brain, uncovering previously undetected activated nuclei (Figure 5A, Supplementary Figure 2). Newly identified activated nuclei in the CTX and the CNU included PL/ILA and ENT, aPVT in the thalamus, TU, ARH, PVi, and PVH in the hypothalamus, and PB in the hindbrain. All nuclei activated by ISO in this study were functionally classified and depicted in Figure 5C. Our results confirmed the activation of several brain regions involved in arousal-related nuclei, such as the prelimbic area (PL) and infralimbic areas (ILA), and the paraventricular nucleus (PVT) [22, 23]. Additionally, we observed activation in previously reported analgesia-related nuclei, including CeA and LC, as well as the parabrachial nucleus (PB) [18, 30, 35]. We also found activation in neuroendocrine function-related nuclei of the hypothalamus, such as TU, PVi, ARH, PVH, and SON. Moreover, we identified activations related to ISO induced side effects, such as in the piriform cortex (PIR) [36] and ENT [37], which may be stimulated by ISO odor, and the solitary tract nucleus (NTS), potentially responsible for ISO induced vomiting [38]. The only activated nucleus in the midbrain was the Edinger-Westphal nucleus (EW). Recent research has found that sevoflurane activates EW and is involved in sleep induction and maintenance of anesthesia, suggesting its crucial role in general anesthesia [39]. By comparing the ISO and KET induced c-Fos expression, we summarized the brain regions activated by both anesthetics in Figure 5D. Despite variations in molecular targets, the coactivation of regions such as PL/ILA, aPVT, CeA, PVH, SON, EW, PB, LC, and NTS by both ISO and KET suggests an overlapping neuronal circuitry that influences sleep-wake regulation, analgesia, and neuroendocrine functions. This shared neural circuitry may potentially offer a common mechanism across the two anesthetics for the maintenance of general anesthesia.

c-Fos expression in distinct brain regions after exposure to ISO.

Earlier version: (A) Representative of brain regions with statistical differences c-Fos+ cells between the ISO group and home cage mice. Scale bar, 200 µm. (B) Cell counts of home cage group and ISO group compared with T-test, P values < 0.05. Data are represented as mean ± SEM (C) Schematic cross-section of the mouse brain showing activated brain regions by ISO. Different colors indicate various functionally relevant nuclei. Red signifies nuclei involved in sleep-wake regulation, blue-green in pain management, blue in neuroendocrine function, pink in side-effect management, and white denotes nuclei exhibiting mixed functionalities. (D) The Venn diagram shows brain regions that are co-activated by ISO and KET and differentially activated brain regions(Manual calibration and T-test correction).

Revised version: (A) Representative of brain regions with statistical differences c-Fos+ cells between the ISO group and home cage mice. Scale bar, 200 µm. (B) Representation of cell counts in brain regions exhibiting significant statistical differences between home cage and ISO. Data are shown as mean ± SEM. P values < 0.05. (C) Schematic cross-section of the mouse brain showing activated brain regions by ISO. Different colors indicate various functionally relevant nuclei. Red signifies nuclei involved in sleep-wake regulation, blue-green in pain management, blue in neuroendocrine function, pink in side-effect management, and white denotes nuclei exhibiting mixed functionalities. (D) The Venn diagram shows brain regions that are co-activated by ISO and KET and differentially activated brain regions.

Network generation and Hub identification

Earlier version: Previous research has established that general anesthesia is mediated by various brain regions [7, 16–18]. c-Fos expression serves as an indicator of neuronal activity, providing a single index of activation per region per animal. By examining the covariance of c-Fos expression across animals within each group, we can infer interactions between brain regions and identify functional networks engaged during general anesthesia [37]. Highly correlated activities across different brain regions are presumed to constitute components of a functional network, reflecting the complex interplay of brain areas during general anesthesia. We first employed c-Fos expression as a neuronal activity marker and calculated a comprehensive set of interregional correlations for four groups. The matrices exhibited interregional correlations for the number of c-Fos-positive cells in each condition (Figure 6A). Network graphs were generated by extracting Pearson’s coefficients > 0.82, as well as significant positive correlations (P < 0.05), from these matrices to construct functional brain networks (Figure 6B). In comparison to the control group, isoflurane (ISO) slightly decreased interconnections between regions (network density: 0.13 vs. 0.10; edges: 267 vs. 198), with no significant difference in mean interregional correlation coefficients between the ISO group and the home cage group (Fisher Z = -0.018, P = 0.98). Conversely, ketamine (KET) significantly increased the network’s connectivity density (network density: 0.13 vs. 0.47; edges: 265 vs. 1008) and showed a significant increase in mean interregional correlation coefficients compared to the saline group (Fisher Z = 3.54, P < 0.001) (Figure 6C). These findings suggest that ISO may exert a mild inhibitory effect on functional network connectivity, whereas KET appears to enhance interregional correlation following administration (Figure 6D).

Revised version: Previous research has established that general anesthesia is mediated by various brain regions [7, 16–18]. c-Fos expression serves as an indicator of neuronal activity, providing a single index of activation per region per animal. By examining the covariance of c-Fos expression across animals within each group, we can infer interactions between brain regions and identify functional networks engaged during general anesthesia [40]. Consequently, we initially incorporated a total of 63 brain regions, including nuclei that were significantly activated by ISO and KET (as indicated by positive findings in Figures 4 and 5) and those previously reported to be associated with the mechanisms of general anesthesia and sleep-wakefulness [41]. Highly correlated activities across different brain regions are presumed to constitute components of a functional network, reflecting the complex interplay of brain areas during general anesthesia. Using c-Fos expression as a marker for neuronal activity and calculated a comprehensive set of interregional correlations for four groups. The matrices exhibited interregional correlations for the number of c-Fos-positive cells in each condition (Figure 6A). Network graphs were generated by extracting Pearson’s coefficients > 0.82, as well as significant positive correlations (P < 0.05), from these matrices to construct functional brain networks (Figure 6B). Compared to the control group, isoflurane (ISO) slightly decreased interconnections between regions (network density: 0.13 vs. 0.10; edges: 267 vs. 198), with no significant difference in mean interregional correlation coefficients between the ISO group and the home cage group (Fisher Z = −0.018, P = 0.98). Conversely, ketamine (KET) significantly increased the network’s connectivity density (network density: 0.13 vs. 0.47; edges: 265 vs. 1008) and showed a significant increase in mean interregional correlation coefficients compared to the saline group (Fisher Z = 3.54, P < 0.001) (Figure 6C). These findings suggest that ISO may exert a mild inhibitory effect on functional network connectivity, whereas KET appears to enhance interregional correlation following administration (Figure 6D).

Generation of anesthetics-induced networks and identification of hub regions.

Earlier version: (A) Matrices showing interregional correlations for c-Fos expression at the home cage, ISO, saline, and KET. Colors reflect correlation strength based on Pearson’s r (color bar, right). Axes are numbered and correspond to brain regions listed in Supplementary Table 5. (B) Network graphs were generated by significant positive correlations (P < 0.05), as well as Pearson’s r > 0.82. The widths of the edges are proportional to the strength of the correlations and the size of the nodes is proportional to the degree of the node. Colors represent the major brain division (red, cerebral nuclei; purple, hypothalamus; orange, midbrain; brown, thalamus; blue, thalamus; green, the cerebral cortex. Network densities were noted in the left. (C) Mean r was calculated from interregional correlation coefficients.

Revised version: (A) Matrices showing interregional correlations for c-Fos expression at the home cage, ISO, saline, and KET. Colors reflect correlation strength based on Pearson’s r (color bar, right). Axes are numbered and correspond to brain regions listed in Supplementary Table 5. (B) Network graphs were generated by significant positive correlations (P < 0.05), as well as Pearson’s r > 0.82. The widths of the edges are proportional to the strength of the correlations and the size of the nodes is proportional to the degree of the node. Colors represent the major brain division (red, cerebral nuclei; purple, hypothalamus; orange, midbrain; brown, thalamus; blue, thalamus; green, the cerebral cortex. Network densities were noted in the left. (C) Mean r was calculated from interregional correlation coefficients. Data were analyzed using one-way ANOVA, followed by Tukey’s post hoc test. ***P < 0.001.

Earlier version: Hubs are nodes that occupy critical central positions within the network, enabling the network to function properly. To determine the brain regions that serve as the hubs for ISO and KET induced functional networks, we calculated each node’s degree (the number of links that connect it) and betweenness centrality (the fraction of all shortest paths in the network that contain a given node). Nodes with a high degree and betweenness centrality typically have a lot of connections. We ranked each node according to its degree and betweenness centrality and extracted nodes with a rank greater than the 80th percentile in its network. Additionally, we segmented the functional network into non-overlapping modules using a spectral community detection algorithm and calculated the within-modal degree Z-score (the degree of nodes within a module, which indicates within-module connectivity) as well as participation coefficients (the distribution of the edges of a node between other modules, which indicates within-module connectivity) [38]. Nodes with relatively high values for both parameters were considered relatively significant within and outside the network module. We classified nodes with participation coefficients > 0.4 and Z scores > 1 as connector hubs [39]. The LC, exhibiting elevated degree and betweenness centrality as well as relatively higher Z-scores and participation coefficients, may play a pivotal role in mediating isoflurane-induced general anesthesia, given its known involvement in arousal, attention, and analgesia. In contrast, the somatosensory cortex (SS) functions as a connector hub in the KET group, indicating its integrative and coordinating role in ketamine-induced dissociative anesthesia (Figure 7D). The LSc served as the central hub in the saline group (Figure 7C). No significant findings were observed in the within-modal degree Z-score and participation coefficient analyses for the home cage group. However, the analysis did reveal a relatively high degree and betweenness centrality of APNs (Figure 7A), suggesting that the within-module connections are relatively independent, or that the pathways of inter-module information transmission do not depend on specific connector nodes.

Revised version: Hubs are nodes that occupy critical central positions within the network, enabling the network to function properly. To determine the brain regions that serve as the hubs for ISO and KET induced functional networks, we calculated each node’s degree (the number of links that connect it) and betweenness centrality (the fraction of all shortest paths in the network that contain a given node). Nodes with high degree and betweenness centrality values typically exhibit many connections. We ranked each node according to its degree and betweenness centrality and extracted nodes with a rank greater than the 80th percentile in its network (Figure 7). Additionally, we segmented the functional network into non-overlapping modules using a spectral community detection algorithm and calculated the within-modal degree Z-score (the degree of nodes within a module, which indicates within-module connectivity) as well as participation coefficients (the distribution of the edges of a node between other modules, which indicates within-module connectivity) [42]. Nodes with relatively high values for both parameters were considered relatively significant within and outside the network module. We classified nodes with participation coefficients > 0.4 and Z scores > 1 as connector hubs [43]. The LC, exhibiting elevated degree and betweenness centrality as well as relatively higher Z-scores and participation coefficients, may play a pivotal role in mediating isoflurane-induced general anesthesia, given its known involvement in arousal, attention, and analgesia. In contrast, the somatosensory cortex (SS) functions as a connector hub in the KET group, indicating its integrative and coordinating role in ketamine-induced dissociative anesthesia (Figure 7D). The LSc served as the central hub in the saline group (Figure 7C). No significant findings were observed in the within-modal degree Z-score and participation coefficient analyses for the home cage group. However, the analysis did reveal a relatively high degree and betweenness centrality of APNs (Figure 7A), suggesting that the within-module connections are relatively independent, or that the pathways of inter-module information transmission do not depend on specific connector nodes.

Identification of hub regions in the home cage (A), ISO (B), NS (C), and KET (D) groups.

Earlier version: The degree and betweenness centrality for each brain region are ranked in descending order. Black columns indicate the overlap of brain regions that rank within the top 20% for both degree and betweenness centrality. The dotted line signifies the ranking threshold, encompassing greater than 80% of the brain regions in our network. The within-community Z-scores and participation coefficients are calculated for each brain region in the networks. Brain regions with a participation coefficient greater than 0.4 and a within-module degree Z-score greater than 1 are defined as hubs, as represented by the red dots.

Revised version: Each brain region is ranked in descending order based on degree and betweenness centrality, respectively. The regions to the left of the dotted line represent the top 20% in our network. Black columns highlight brain regions that rank similarly in both degree and betweenness centrality. The within-module Z-scores and participation coefficients were calculated for each brain region in the networks. Brain regions characterized by a participation coefficient greater than 0.4 and a within-module degree Z-score greater than 1 are defined as hubs and are denoted by red dots.

The possible framework for ketamine and isoflurane induced unconsciousness.

Earlier version: The distinct pathways of ketamine and isoflurane induced unconsciousness can be explained by two contrasting mechanisms. The “top-down” process attributes KET’s effect to widespread cortical activation (represented in yellow), with the somatosensory cortex (SS) acting as the central node in the functional network (depicted in blue) and relative inhibition of hypothalamic sleep-promoting nuclei. Conversely, the “bottom-up” approach posits that isoflurane induced unconsciousness arises from the activation of hypothalamic sleep-promoting regions (indicated in yellow) and relative inhibition of cortical and thalamic nuclei, with the locus coeruleus (LC) serving as the hub node in the isoflurane induced functional network. Nuclei activated by both anesthetics are shown in green. Adapted from [3, 4, 47]. PL, prelimbic area; ILA, infralimbic areas; SON, supraoptic nucleus; PVH, paraventricular hypothalamic nucleus; TU, tuberal nucleus; LC, locus coeruleus; SS, somatosensory cortex; CTX: cortex; TH: thalamus; HY, hypothalamus; MB; mid-brain; HB, hindbrain.

Revised version: The distinct pathways of ketamine and isoflurane induced unconsciousness can be explained by two contrasting mechanisms. The “top-down” process attributes KET’s effect to widespread cortical activation (represented in yellow), with the somatosensory cortex (SS) acting as the central node in the functional network (depicted in blue) and relative inhibition of hypothalamic sleep-promoting nuclei. Conversely, the “bottom-up” approach posits that isoflurane induced unconsciousness arises from the activation of hypothalamic sleep-promoting regions (indicated in yellow) and relative inhibition of cortical and thalamic nuclei, with the locus coeruleus (LC) serving as the hub node in the isoflurane induced functional network. Nuclei activated by both anesthetics are shown in green. Adapted from [3, 4, 52]. PL, prelimbic area; ILA, infralimbic areas; SON, supraoptic nucleus; PVH, paraventricular hypothalamic nucleus; TU, tuberal nucleus; LC, locus coeruleus; SS, somatosensory cortex; CTX: cortex; TH: thalamus; HY, hypothalamus; MB; midbrain; HB, hindbrain.

Discussion

In this study, we conducted a comparative analysis of the effects of two general anesthetics, isoflurane (ISO) and ketamine (KET), on c-Fos expression throughout the brain. By employing principal component analysis (PCA), we were able to thoroughly examine the responses of individual brain regions to each anesthetic agent. Our findings reveal that KET dominantly activates the cerebral cortex yet suppresses subcortical regions, reflecting a top down mechanism of action, while ISO predominantly stimulates subcortical brain regions with relative cortical inhibition, substantiating its bottom up mechanism of action [3, 4]. Further functional analysis of brain networks, based on c-Fos expression, identified the somatosensory cortex (SS) and the locus coeruleus (LC) as central nodes for KET and ISO, respectively, highlighting the crucial roles of LC and SS under ISO and KET induced unconsciousness.

Earlier version: Our results demonstrate that ISO activates the sleep-promoting VLPO nucleus, the aPVT-infralimbic loop [20], and broadly inhibits the cortex, supporting a bottom-up mechanism of ISO induced unconsciousness. Nonetheless, our findings also reveal the activation of arousal related nuclei, such as PB and LC, which implies that the influence of isoflurane on consciousness may not solely rely on suppressing arousal centers, but rather through a more intricate relationship than formerly recognized. Identifying cell types and their dynamic changes during anesthesia will be crucial for clarifying their role in ISO induced unconsciousness.

Revised version: Our results demonstrate that ISO activates the sleep-promoting VLPO nucleus, the aPVT-infralimbic loop [23], and broadly inhibits the cortex, supporting a bottom-up mechanism of ISO induced unconsciousness. Nonetheless, our findings also reveal the activation of arousal related nuclei, such as PB and LC, which implies that the influence of isoflurane on consciousness may not solely rely on suppressing arousal centers, but rather through a more intricate relationship than formerly recognized. Identifying cell types and their dynamic changes during anesthesia will be crucial for clarifying their role in ISO induced unconsciousness.

Earlier version: While c-Fos-based functional network analysis offers lower temporal resolution, its single-cell resolution across the entire brain allows for the inclusion of midbrain and hindbrain regions, supplementing previous fMRI analyses. Our observations reveal that ISO mildly inhibits network density, and through graph theoretical analysis, we identify the LC as a highly connected hub, high-lighting the critical role of the brainstem in ISO induced general anesthesia. The LC performs a wide range of functions in mice, including arousal, pain modulation, attention, stress response, and neuroprotection. Studies have shown that chemical activation of the LC increases whole-brain functional connectivity, attributed to its role as the primary source of norepinephrine (NE) and its extensive influence on nearly the entire brain [40]. The significant LC activation and its central position within the functional network underlying ISO induced unconsciousness suggest that the LC plays a crucial part in maintaining and integrating the entire unconsciousness functional network, emphasizing the involvement of LC in the bottom-up paradigm of ISO induced unconsciousness.

Revised version: While c-Fos-based functional network analysis offers lower temporal resolution, its single-cell resolution across the entire brain allows for the inclusion of midbrain and hindbrain regions, supplementing previous fMRI analyses. Our observations reveal that ISO mildly inhibits network density, and through graph theoretical analysis, we identify the LC as a highly connected hub, highlighting the critical role of the brainstem in ISO induced general anesthesia. The LC performs a wide range of functions in mice, including arousal, pain modulation, attention, stress response, and neuroprotection. Studies have shown that chemical activation of the LC increases whole-brain functional connectivity, attributed to its role as the primary source of norepinephrine (NE) and its extensive influence on nearly the entire brain [44]. The significant LC activation and its central position within the functional network underlying ISO induced unconsciousness suggest that the LC plays a crucial part in maintaining and integrating the entire unconsciousness functional network, emphasizing the involvement of LC in the bottom-up paradigm of ISO induced unconsciousness.

Earlier version: Mashour et al. proposed that anesthesia-induced unconsciousness encompasses not only the modulation of lower-level brain activity but also top down neural processing [3, 4]. Within this top down framework, anesthetics diminish consciousness by interfering with cortical and thalamocortical circuits responsible for neural information integration. Our study discovered that KET administration substantially activated cortical and subcortical arousal-promoting nuclei while concurrently causing relative thalamic suppression, with only the RE and TRS exhibiting activation. This suggests that thalamic inhibition may lead to a reduction in thalamocortical communication, which is characterized by the inability to perceive the external environment and results in disconnection from reality. Graph theoretical analysis also identified the somatosensory cortex (SS) as the hub node of the KET induced functional network. As a critical cortical area, SS is responsible for sensory processing, motor control, and cognitive functions [41]. Previous studies have demonstrated that local KET administration to SS recapitulates the effects of systemic KET on both the switch in pyramidal cell activity and dissociative-like behavior, implying that SS may serve as a key target for KET induced dissociation [42]. Our findings that SS acts as a hub node suggest that KET may modulate brain network function by influencing connectivity between SS and other brain regions, thereby affecting the behavior and cognitive states of mice. This further supports the significance of cortical areas during KET anesthesia.

Revised version: Mashour et al. proposed that anesthesia-induced unconsciousness encompasses not only the modulation of lower-level brain activity but also top down neural processing [3, 4]. Within this top down framework, anesthetics diminish consciousness by interfering with cortical and thalamocortical circuits responsible for neural information integration. Our study discovered that KET administration substantially activated cortical and subcortical arousal-promoting nuclei while concurrently causing relative thalamic suppression, with only the RE and TRS exhibiting activation. This suggests that thalamic inhibition may lead to a reduction in thalamocortical communication, which is characterized by the inability to perceive the external environment and results in disconnection from reality. Graph theoretical analysis also identified the somatosensory cortex (SS) as the hub node of the KET induced functional network. As a critical cortical area, SS is responsible for sensory processing, motor control, and cognitive functions [45]. Previous studies have demonstrated that local KET administration to SS recapitulates the effects of systemic KET on both the switch in pyramidal cell activity and dissociative-like behavior, implying that SS may serve as a key target for KET induced dissociation [46]. Our findings that SS acts as a hub node suggest that KET may modulate brain network function by influencing connectivity between SS and other brain regions, thereby affecting the behavior and cognitive states of mice. This further supports the significance of cortical areas during KET anesthesia.

Identifying shared neural features between KET and ISO is essential for understanding anesthetic-induced unconsciousness. The coactivation of sleep-wake regulation-related regions, such as PL/ILA and aPVT, along with analgesia-related nuclei like CeA, PB, and LC, suggests a shared mechanism for sleep-wake regulation and the common pathways for pain relief. This observation provides valuable insights into the fundamental mechanisms of anesthesia-induced hypnosis and analgesia. Additionally, the coactivation of neuroendocrine-related nuclei, including PVH and SO in the hypothalamus, raises questions about the potential influence of anesthetics on hormonal release and homeostatic regulation. Other coactivated nuclei, such as EW and NTS, warrant further investigation of their roles in anesthesia. In summary, the coactivated nuclei imply a potential shared neuronal circuitry for general anesthesia, encompassing common features like unconsciousness, analgesia, and autonomic regulation, regardless of the specific molecular targets of each drug. Future research could examine coactivated brain regions by the two anesthetics or manipulate identified hub nodes to further understand the mechanisms of general anesthesia. In summary, our study reveals distinct and shared neural mechanisms underlying isoflurane and ketamine anesthesia using c-Fos staining and network analysis. Our findings support “top-down” and “bottom-up” paradigms, and the identification of hub nodes and coactivated brain regions suggests shared neurocircuitry for general anesthesia, providing insights into the mechanisms underlying anesthetic-induced unconsciousness and analgesia.

Methods

Animals

All animal experiments were conducted in accordance with the National Institutes of Health guidelines and were approved by the Chinese Academy of Sciences’ Institute of Neuroscience. Adult male wild-type (WT) mice (C57BL/6J) (8-10 weeks old, weight from 22 to 26 g) were purchased from institute-approved vendors (LingChang Experiment Animal Co., China). Mice were individually housed and maintained on a 12 h:12 h light/dark cycle (lights on at 07:00 a.m. and off at 07:00 p.m.) with food and water available ad libitum.

Drug administration

Earlier version: All experiments occurred between 13:00-14:30 (ZT6–ZT7.5). We adapted mice to handling and the anesthesia chamber (10×15×15 cm) for several days to minimize experimental confound-induced c-Fos induction. Adult male mice were handled for the KET group for 10 min per day with normal saline (NS) injected intraperitoneally (i.p.) for three consecutive days at 13:00. On day five, a randomly chosen mouse received an injection of Ketamine (Gutian Medicine, H35020148), and the control groups (n=8) received the same volume of saline. ISO group (n=6) mice were handled and inhaled 1.5% isoflurane (RWD Life Science, 1903715) at 13:00 on day four in the chamber. Meanwhile, the control groups (n=6) were left undisturbed in their home cages prior to sampling. We confirmed the loss of righting reflex at 5 min after anesthetics exposure. For 90 min after KET injection or ISO inhalation, mice were deeply anesthetized with 5% ISO and transcardially perfused with 35ml 0.1 M phosphate-buffered saline (PBS) followed by 35ml 4% paraformaldehyde (PFA). The brains were then removed and postfixed overnight with 4% PFA. Following fixation, the brains were dehydrated for 48 hours with 30% sucrose (wt/vol) in PBS. Coronal sections (50 µm) of the whole brain were cut using a cryostat (HM525 NX, Thermo Scientific) after being embedded with OCT compound (NEG-50, Thermo Scientific) and freezing.

Revised version: All experiments occurred between 13:00-14:30 (ZT6–ZT7.5). We adapted mice to handling and the anesthesia chamber (10×15×15 cm) for several days to minimize experimental confound-induced c-Fos induction. Adult male mice were handled for the KET group for 10 min per day with normal saline (NS) injected intraperitoneally (i.p.) for three consecutive days at 13:00. On day five, a randomly chosen mouse received an injection of Ketamine (Gutian Medicine, H35020148), and the control groups (n=8) received the same volume of saline. ISO group (n=6) mice were handled and inhaled 1.5% isoflurane (RWD Life Science, 1903715) at 13:00 on day four in the chamber. Meanwhile, the control groups (n=6) were left undisturbed in their home cages prior to sampling. We confirmed the loss of righting reflex at 5 min after anesthetics exposure. For 90 min after KET injection or ISO inhalation, mice were deeply anesthetized with 5% ISO and transcardially perfused with 35 ml 0.1 M phosphate-buffered saline (PBS) followed by 35 ml 4% paraformaldehyde (PFA). The brains were then removed and postfixed overnight with 4% PFA. Following fixation, the brains were dehydrated for 48 hours with 30% sucrose (wt/vol) in PBS. Coronal sections (50 µm) of the whole brain were cut using a cryostat (HM525 NX, Thermo Scientific) after being embedded with OCT compound (NEG-50, Thermo Scientific) and freezing.

Added in the revised version: Anesthesia depth measurement

Added in the revised version: EEG and EMG recordings were conducted as described previously[20]. Specifically, a stainless steel electrode was positioned over the frontal cortex for EEG measurements, while an insulated EMG electrode was inserted into the neck muscles. Additionally, a reference electrode was placed above the cerebellum. The EEG/EMG data were captured using TDT system-3 amplifiers, specifically RZ5 + RA16PA and RZ2 + PZ5 configurations. To measure the depth of anesthesia using EEG/EMG signals, a fast Fourier transform spectral analysis was employed, featuring a frequency resolution of 0.18 Hz. Data was noted every 5 seconds through a MATLAB tool and subsequently verified manually by experts. The assessment was based on EMG power and the ratio of EEG’s δ power (0.5-4Hz) to θ power (6-10Hz). EEG and EMG power values within 30 minutes post-administration were normalized to a 5-minute pre-administration baseline.

Immunohistochemistry

Earlier version: One out of every three brain slices (100 µm intervals) of each whole brain was washed three times with 0.1 M phosphate-buffered saline (PBS) for 10 min and permeabilized for 30 minutes at room temperature with 0.3% Triton X-100 in PBS (PBST). Slices were incubated for 2 hours at room temperature with 2% normal donkey serum (Sigma, G6767) in PBS overnight at 4°C with c-Fos primary antibodies (226003, Synaptic Systems; 1:500) diluted in PBS with 1% donkey serum. After three washes with PBST, slices were incubated with the Cy3 donkey anti-rabbit (711165152, Jackson; 1:200) secondary antibody for 2 hours at room temperature. Immunostained slices were mounted with VECTASHIELD mounting medium with DAPI and then scanned under a fluorescent microscope equipped with a 10× objective (VS120, Olympus) or a confocal microscope with a 20× objective (FV300, Olympus).

Revised version: One out of every three brain slices (100 µm intervals) of each whole brain was washed three times with 0.1 M phosphate-buffered saline (PBS) for 10 min and permeabilized for 30 minutes at room temperature with 0.3% Triton X-100 in PBS (PBST). Slices were incubated for 2 hours at room temperature with 2% normal donkey serum (Sigma, G6767) in PBS overnight at 4°C with c-Fos primary antibodies (226003, Synaptic Systems; 1:500) diluted in PBS with 1% donkey serum. After three washes with PBST, slices were incubated with the Cy3 donkey anti-rabbit (711165152, Jackson; 1:200) secondary antibody for 2 hours at room temperature. Immunostained slices were mounted with VECTASHIELD mounting medium with DAPI and then scanned under a fluorescent microscope equipped with a 10× objective (VS120, Olympus) or a confocal microscope with a 20× objective (FV300, Olympus). Due to the limited scope of a single field of view, either 2 or 4 adjacent fields of view were stitched together to offer a comprehensive representation of specific brain regions, including the PL, ILA, LSc, LSr, SS, and VISC.

Quantification of c-Fos positive cells

Earlier version: The procedures used for c-Fos analysis were based on previously research [19]. A custom-written software package was employed for the analysis of brain images. The software consists of four modules: atlas rotation, image registration, signal detection, and quantification/visualization.

Revised version: The procedures used for c-Fos analysis were based on previous research [19]. A custom-written software package was employed for cellular quantification within brain images. The software consists of three components: atlas rotation, image registration, and signal detection.

Earlier version: Rotation module

Revised version: Atlas rotation module

Added in the revised version: Utilizing the Allen Mouse Brain Atlas, this module allows rotations in both horizontal and vertical dimensions to align with mouse brain samples. To determine the appropriate rotation angles, we manually pinpointed anatomical landmarks corresponding to the samples. For the left-right axis rotation, we chose points from the CA3 pyramidal layer in both hemispheres and the posterior slice of the dentate gyrus, specifically the granule cell layer. For the dorsal-ventral axis rotation, we identified key anatomical landmarks. These include the initial connections of the anterior commissure and corpus callosum between the hemispheres, as well as ventral regions like the interpeduncular nucleus and the suprachiasmatic nucleus. After determining these rotation angles, we adjusted the reference atlas to match our samples.

Removed in the revised version: The rotation module allows the Allen Mouse Brain Atlas to be rotated in three dimensions at arbitrary angles to fit individual samples. Anatomical landmarks were manually selected to estimate rotation angles. The most posterior slices containing the CA3 in each hemisphere were used to calculate the rotation angle about the left-right axis. The most anterior slice with the anterior commissure crossing the midline and the most posterior slice with the corpus callosum crossing the midline were used to calculate the rotation angle about the dorsal-ventral axis. The reference atlas was then rotated using these estimated angles to match the experimental brain’s sectioning angle.

Earlier version: Registration module

Revised version: Image registration module

Added in the revised version: This module uses a tool that aligns brain slice images with a reference atlas, facilitating the alignment of overall brain region distribution. Registration starts by matching the coronal plane of the sample sections to the atlas. After defining the external boundaries of the brain section, the system performs geometric transformations on the section to optimize its fit with the atlas.

Removed in the revised version: The registration module is an image alignment software that uses reference points to align brain section images with a rotated 3D reference atlas for further quantification. First, reference points in both the atlas and the brain image were chosen. The module then applied geometric transformations to the brain section to optimize the match between the brain image and the atlas’s reference points. After the transformation, the image was compared to the atlas, and any necessary adjustments were made manually.

Earlier version: Detection module

Revised version: Signal detection module

Added in the revised version: The detection module is designed to automatically label c-Fos+ cells. Post-detection, a manual verification was performed on each digitized brain section image to ensure the accuracy and precision of the c-Fos+ cell markings.

Removed in the revised version: The detection module manually counts the position of c-Fos positive cells in each digitized brain section image.

Principal components analysis (PCA) of the activity patterns at whole brain

Earlier version: We employed arithmetic to calculate the cell density in each brain region: Nr: The number of c-Fos+ cells in each brain region. Vr: The volume of each brain region. The density of c-Fos from 201 brain regions were concatenated for each mouse (Home cage, n = 6; ISO, n = 6; Saline, n = 8; KET, n = 6). Principal components analysis was performed on the concatenated matrix containing data from four conditions, using singular value decomposition. We then selected the first two principal components (PCs), which accounted for 56.6% of the variance in Figure 2A and 44.5% in Figure 3A.

Revised version: We utilized arithmetic methods to compute a logarithm ratio of the c-Fos density in a specific brain region for PCA analysis. Specifically, the cell density of a single brain region (defined as the number of c-Fos positive cells in a brain region divided by the volume of that region) was divided by the global c-Fos cell density (calculated as the total number of c-Fos positive cells throughout the entire brain divided by the total brain volume). Subsequently, the logarithm of this ratio value was taken. The formula is as follows: Nr: The number of c-Fos+ cells in each brain region. Vr: The volume of each brain region. Principal components analysis was performed on the concatenated matrix containing data from four conditions, using singular value decomposition. We then selected the first two principal components (PCs), which accounted for 56.6% of the variance in Figure 2A and 44.5% in Figure 3A.

Network generation

Earlier version: To evaluate how functional connectivity changed under general anesthetics in WT mice, we extracted 63 brain regions from major brain subdivisions (cerebral cortex, cerebral nuclei, thalamus, hypothalamus, midbrain, and hindbrain) listed in Supplementary Table 2. Correlation matrices were generated by computing Pearson correlation coefficients from interregional c-Fos expression in the 63 regions. Mean correlations were calculated to assess changes in functional connectivity between these major subdivisions of the brain. Weighted undirected networks were constructed by considering correlations with Pearson’s r ≥0.82, corresponding to a one-tailed significance level of P<0.05 (uncorrected). The nodes in the networks represent brain regions, and the correlations that survived thresholding were considered connections. Theoretical graph analysis was performed using Brain Connectivity Toolbox (https://sites.google.com/site/bctnet/, version 2019-03-03) in MATLAB R2021 (The MathWorks Inc.) [43]. Network visualization was performed using Cytoscape (version 3.2.1) [44].

Revised version: To evaluate how functional connectivity changed under general anesthetics in WT mice, we extracted 63 brain regions from major brain subdivisions (cerebral cortex, cerebral nuclei, thalamus, hypothalamus, midbrain, and hindbrain) listed in Supplementary Table 2. Correlation matrices were generated by computing Pearson correlation coefficients from interregional c-Fos expression in the 63 regions. Mean correlations were calculated to assess changes in functional connectivity between these major subdivisions of the brain. Weighted undirected networks were constructed by considering correlations with Pearson’s r ≥0.82, corresponding to a one-tailed significance level of P<0.05 (uncorrected). The nodes in the networks represent brain regions, and the correlations that survived thresholding were considered connections. Theoretical graph analysis was performed using Brain Connectivity Toolbox (https://sites.google.com/site/bctnet/, version 2019-03-03) in MATLAB R2021 (he MathWorks Inc.) [47]. Network visualization was performed using Cytoscape (version 3.2.1) [48].

Hub identification

Earlier version: Network centrality was evaluated using degree, betweenness, within-module z-scores (representing within-module connectivity), and participation coefficient (indicating between-module connectivity). These measures were computed for all nodes to identify potential hub regions [45]. Modularity was assessed employing Newman’s spectral community detection algorithm [46]. Degree represents the number of edges connected to a node, while betweenness denotes the number of shortest paths traversing a given node. Nodes with elevated betweenness centrality are involved in numerous shortest paths.

Revised version: Network centrality was evaluated using degree, betweenness, within-module z-scores (representing within-module connectivity), and participation coefficient (indicating between-module connectivity). These measures were computed for all nodes to identify potential hub regions [49]. Modularity was assessed employing Newman’s spectral community detection algorithm [50]. Degree represents the number of edges connected to a node, while betweenness denotes the number of shortest paths traversing a given node. Nodes with elevated betweenness centrality are involved in numerous shortest paths.

Statistical analysis

Earlier version: The sample size was determined based on prior studies [13, 14]. Normality of data distribution was assessed using the Shapiro-Wilk W test. Unpaired Student’s t-test was applied for normally distributed data, while the Mann-Whitney U test was employed for non-normally distributed data. Data are presented as mean ± SEM, with all statistical tests being two-sided. Pearson correlation coefficients (R) were transformed into Z-scores using Fisher’s Z transformation before computing group means and conducting statistical comparisons. GraphPad Prism 9.0 (GraphPad Software, USA) and MATLAB R2021 (Mathworks Inc.) were utilized for statistical analyses. A P-value less than 0.05 was considered statistically significant.

Revised version: Sample size was determined based on prior studies [13, 14]. Data are presented as mean ± SEM, and all statistical tests were two-sided. For the c-Fos cell percentages in Figure 2C, we employed the Mann-Whitney test with a false discovery rate correction, setting a Q-value at 0.05. In Figures 4 and 5, differences in c-Fos cell counts for both KET and ISO were evaluated using an independent t-test. P-values were adjusted using the procedure by Benjamini, Krieger, and Yekutieli with Q set at 0.05 [51]. Figure 6C and Supplementary Figure 1 utilized a one-way ANOVA followed by a post-hoc Tukey correction. Pearson correlation coefficients (R) were transformed into Z-scores using the Fisher’s Z transformation before computing group means and making statistical comparisons. All statistical analyses were conducted using GraphPad Prism 9.0 (GraphPad Software, USA) and MATLAB R2021 (Mathworks Inc.). A P-value of less than 0.05 was deemed statistically significant.

Supporting information